For the meeting today, we can treat this as the template notebook we want to provide. Later, I will update it to contain good standards for grading cells and setup, a lot of which is from https://nbgrader.readthedocs.io/en/stable/configuration/student_version.html
it looks like all cells are run and then tests are matched?

Also give info on how to set cell types, etc., cell id conventions, ...

- Why can I not see the file tree when I edit assignment
- We definitely do not want them to have access to the terminal, right (?)
    -- In CMU-INAI, for example, we had to issue a bugfix by asking them to download a package through the terminal





<!-- grade_id: gcd_answer, points: 0, grade: false, solution: true -->
```python
# Here is a test question 1

def GCD(a, b):
    # Return the GCD of a and b using the euclidian algorithm
    ### BEGIN SOLUTION
    # This a simple recursive solution
    return GCD(b % a, a) if a else b
    ### END SOLUTION
```

<!-- grade_id: gcd_tests, points: 10, grade: true, solution: false -->
```python
# Here are tests for question 1
assert GCD(60, 45) == GCD(45, 60) == 15
assert GCD(200, 1) == 1
assert GCD(25, 65) == 5
### BEGIN HIDDEN TESTS
# It can be helpful to redefine the function in as given in the solution, but with a different name, to compare values.
def sol_GCD(a, b):
    return GCD(b % a, a) if a else b

# Correctness tests
for x in range(1, 100):
    for y in range(1, 50):
        assert GCD(x, y) == sol_GCD(x, y), f"GCD failed for GCD({x}, {y})"
        
# Run efficiency tests?

### END HIDDEN TESTS

```

```python

```