# jupyter-notebook Assignment Template

This file is a template that contains basic structure and usage for creating an nbgrader-compatible notebook assignment.

The actual format is pretty flexible. In this guide, we provide an example question and show what nbgrader will do to the notebook when creating the student assignment version. You can use these principles to format the document as you wish, but it is reccommended to stick to this format for simplicity and consistency. For more information on how nbgrader works in detail, see the documentation at https://nbgrader.readthedocs.io/en/stable/configuration/student_version.html.

Here is an example question.

 ### Question 1: The Euclidian Algorithm
 
 Write the `GCD` function below, which takes two numbers and returns their GCD. You should refer to the lecture notes on how to implement the Euclidian algorithm when completing the function.

<!-- grade_id: cell-99ccc0ea01c8e88d, points: 0, grade: false, solution: true -->
```python
def GCD(a, b):
    # Return the GCD of a and b using the euclidian algorithm
    ### BEGIN SOLUTION
    # This a simple recursive solution
    return GCD(b % a, a) if a else b
    ### END SOLUTION
```

<!-- grade_id: cell-d75ad048d692ef5a, points: 1, grade: true, solution: false -->
```python
assert GCD(60, 45) == GCD(45, 60) == 15
assert GCD(200, 1) == 1
assert GCD(25, 65) == 5

### BEGIN HIDDEN TESTS
def sol_GCD(a, b):
    return GCD(b % a, a) if a else b

# Correctness tests
for x in range(1, 100):
    for y in range(1, 50):
        assert GCD(x, y) == sol_GCD(x, y), f"GCD failed for GCD({x}, {y})"

### END HIDDEN TESTS
```

Let's break this down.

First, we have a markdown cell which introduces the student to the question, including any information or context that may help them. This improves readability for the student.

Next, we provide a cell which has the area that the student will actually write their code. This includes a function we want them to complete, with a brief description of what the function does. Within this section, you will notice `### BEGIN SOLUTION` and `### END SOLUTION` delimiters. When nbgrader creates the student version of the assignment, it will remove everything in these bounds and replace it with (by default) 

`# YOUR CODE HERE`

`raise NotImplementedError`

At the correct indentation level. This allows space for the students to write their solution.

Lastly, there is a cell containing test cases to assess the correctness of the student's code. These cells are run by nbgrader against the student's solution to determine their grade. Within this section, you will notice `### BEGIN HIDDEN TESTS` and `### END HIDDEN TESTS` delimiters. Similarly to before, nbgrader will remove all code in this space when creating the student assignment, but it will not replace it with anything. Therefore, the only tests which the students will actually see are those which are outside of this area. Generally, it is reccommended to provide a few test cases visible to the student so that they may better understand the question.

The tests are in the form `assert <test>, <error_message>`, which will throw an error to nbgrader if the `<test>` is false, providing the `<error_message>` in the grade report.

Your notebook can also contain other markdown cells to direct the student or code cells to help the assignment function.

Once you complete the assignment file, you will have to specify which cells are which type. If you are working in a Codio project, you can do this from the Jupyter toolbar via View->Cell Toolbar->Create Assignment which will show a cell type dropdown for each cell. For more information on how to set these, see the Developing Assignments in Codio guide.